```python
# Initialize Otter
import otter
grader = otter.Notebook("hw02.ipynb")
```

# Homework 2: Names and Tables 

Please complete this notebook by filling in the cells provided. Before you begin, execute the previous cell to load the provided tests.

```python
# Run this cell and ignore the output. It is checking to see if you ran the first cell! 
# If you get a NameError, please run the cell at the very TOP of this notebook!

grader
```

**Helpful Resource:**
- [Python Reference](https://pages.mtu.edu/~lebrown/data1202-s24/reference/index.html): Cheat sheet of helpful array & table methods used in DATA 1202!

**Recommended Readings:**
- [What is Data Science?](http://www.inferentialthinking.com/chapters/01/what-is-data-science.html)
- [Causality and Experiments](http://www.inferentialthinking.com/chapters/02/causality-and-experiments.html) 
- [Programming in Python](http://www.inferentialthinking.com/chapters/03/programming-in-python.html)

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, **please be sure to not re-assign variables throughout the notebook!** For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!


**Note: This homework has hidden tests on it. That means even though tests may say 100% passed, it doesn't mean your final grade will be 100%. We will be running more tests for correctness once everyone turns in the homework.**

Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. Refer to the policies page to learn more about how to learn cooperatively.

You should start early so that you have time to get help if you're stuck.

```python
# Run this cell to set up the notebook, but please don't change it.

import numpy as np
from datascience import *
import d8error
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

# 1. Names and Assignment Statements

**Question 1.1.** When you run the following cell, Python produces a cryptic error message.

```python
4 = 2 + 2
```

Choose the best explanation of what's wrong with the code, and then assign 1, 2, 3, or 4 to `names_q1` below to indicate your answer. **(4 Points)**

1. Python is smart and already knows `4 = 2 + 2`.

2. In Python, it's a rule that the `=` sign must have a variable name to its left, and `4` isn't a variable name.

3. It should be `2 + 2 = 4`.

4. I don't get an error message. This is a trick question.

```python
names_q1 = ...
```

```python
grader.check("q1_1")
```

**Question 1.2.** When you run the following cell, Python will produce another cryptic error message.

```python
two = 3
six = two plus two
```

Choose the best explanation of what's wrong with the code and assign 1, 2, 3, or 4 to `names_q2` below to indicate your answer. **(4 Points)**

1. The `plus` operation only applies to numbers, not the word "two".

2. The name "two" cannot be assigned to the number 3.

3. Two plus two is four, not six.

4. The name `plus` isn't a built-in operator; instead, addition uses `+`.

```python
names_q2 = ...
```

```python
grader.check("q1_2")
```

**Question 3.** Run the following cell.

```python
x = 2
y = 3 * x
x = 4
```

What is `y` after running this cell, and why? Choose the best explanation and assign 1, 2, 3, or 4 to `names_q3` below to indicate your answer. **(4 Points)**

1. `y` is equal to 6, because the second `x = 4` has no effect since `x` was already defined.

2. `y` is equal to 6, because `x` was 2 when `y` was assigned, and 3 * 2 is 6.

3. `y` is equal to 12, because `x` is 4 and 3 * 4 is 12.

4. `y` is equal to 12, because assigning `x` to 4 will update `y` to 12 since `y` was defined in terms of `x`.

```python
names_q3 = ...
```

```python
grader.check("q1_3")
```

# 2. Different Between Majors 

Michigan Tech's Office of Institutional Research provides data on numerous aspects of the campus.  Adapted from the IR website, the table below displays the number of degree recipients in three majors in the 2013-14 and 2021-22 academic years. 

| Major                              | 2013-2014    | 2021-2022   |
|------------------------------------|--------------|-------------|
| Computer Science                   |      34      |    81       |
| Civil Engineering                  |      108     |    77       |
| Computer Engineering               |      35      |    48       |

**Question 1.** Suppose you want to find the **biggest** absolute difference between the number of degree recipients in the two years, among the three majors.

In the cell below, compute this value and call it `biggest_change`. Use a single expression (a single line of code) to compute the answer. Let Python perform all the arithmetic (like subtracting 34 from 81) rather than simplifying the expression yourself. The built-in `abs` function takes a numerical input and returns the absolute value. The built-in `max` function can take in 3 arguments and returns the maximum of the three numbers. **(5 Points)**

```python
biggest_change = ...
biggest_change
```

```python
grader.check("q2_1")
```

**Question 2.2.** Which of the three majors had the **smallest** absolute difference? Assign `smallest_change_major` to 1, 2, or 3 where each number corresponds to the following major:

1. Computer Science  
2. Civil Engineering  
3. Computer Engineering

Choose the number that corresponds to the major with the smallest absolute difference.

You should be able to answer by rough mental arithmetic, without having to calculate the exact value for each major. **(4 Points)** 

```python
smallest_change_major = ...
smallest_change_major
```

```python
grader.check("q2_2")
```

**Question 2.3.**  For each major, define the “relative change” to be the following: $\large{\frac{\text{absolute difference}}{\text{value in 2013-2014}} * 100}$ 

Fill in the code below such that `compsci_relative_change`, `civeng_relative_change` and `compeng_relative_change` are assigned to the relative changes for their respective majors. **(6 Points)**

```python
# Next, define a circumference function.
compsci_relative_change = (abs(...) / 34) * 100 
civeng_relative_change = ...  
compeng_relative_change = ...
print("Relative Change CS:      ", compsci_relative_change)
print("Relative Change CivEng:  ", civeng_relative_change)
print("Relative Change CompEng: ", compeng_relative_change)
```

```python
grader.check("q2_3")
```

**Question 2.4.** Assign `biggest_rel_change_major` to 1, 2, or 3 where each number corresponds to to the following: 

1. Computer Science  
2. Civil Engineering  
3. Computer Engineering

Choose the number that corresponds to the major with the biggest relative change. **(4 Points)**

```python
biggest_rel_change_major = ...
biggest_rel_change_major
```

```python
grader.check("q2_4")
```

# 3. Tables 

Suppose you have 4 apples, 3 oranges, and 3 pineapples.  Below is the code we can use to create a table that contains this information.  It has two columns: `fruit` and `count`. 

```python
fruits = Table().with_columns( 
        "fruit name", ["apple", "orange", "pineapple"], 
        "count",      [4, 3, 3]) 
fruits
```

**Question 3.1.** The file `inventory.csv` contains information about the inventory at a fruit stand.  Each row represents the contents of one box of fruit. Load it as a table named `inventory` using the `Table.read_table()` function. `Table.read_table(...)` takes one argument (data file name in string format) and returns a table. **(4 Points)**


```python
inventory = ...
inventory
```

```python
grader.check("q3_1")
```

**Question 3.2.** Does each box at the fruit stand contain a different fruit? Set `all_different` to `True` if each box contains a different fruit or to `False` if multiple boxes contain the same fruit. **(4 Points)**

*Hint:* You don't have to write code to calculate the True/False value for `all_different`. Just look at the `inventory` table and assign `all_different` to either `True` or `False` according to what you can see from the table in answering the question.


```python
all_different = ...
all_different
```

```python
grader.check("q3_2")
```

**Question 3.3.** The file `sales.csv` contains the number of fruit sold from each box last Saturday.  It has an extra column called `price per fruit ($)` that's the price *per item of fruit* for fruit in that box.  The rows are in the same order as the `inventory` table.  Load these data into a table called `sales`. **(5 Points)**

```python
sales = ...
sales
```

```python
grader.check("q3_3")
```

**Question 3.4.** How many fruits did the store sell in total on that day? **(5 Points) Challenge Problem**

Remember if you have questions on Table functions, please look at the [reference site](https://pages.mtu.edu/~lebrown/data1202-s24/reference/index.html)

```python
total_fruits_sold = ...
total_fruits_sold
```

```python
grader.check("q3_4")
```

You're done with Homework 2!  

**Important submission steps:** 
1. Run the tests and verify that they all pass.
2. Choose **Save Notebook** from the **File** menu, then **run the final cell**. 
3. Click the link to download the zip file.
4. Then submit the zip file to the corresponding assignment on Gradescope. 

**It is your responsibility to make sure your work is saved before running the last cell.**


```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: [1])
```
>>> 1 <= names_q1 <= 4
True
```


Tests for q1_2 (points: [1])
```
>>> 1 <= names_q2 <= 4
True
```


Tests for q1_3 (points: [1])
```
>>> 1 <= names_q3 <= 4
True
```


Tests for q2_1 (points: [1])
```
>>> isinstance(biggest_change, (int, float))
True
```


Tests for q2_2 (points: [1])
```
>>> 1 <= smallest_change_major <= 3
True
```


Tests for q2_3 (points: [0, 0, 0])
```
>>> isinstance(compsci_relative_change, (int, float))
True
```
```
>>> isinstance(civeng_relative_change, (int, float))
True
```
```
>>> isinstance(compeng_relative_change, (int, float))
True
```


Tests for q2_4 (points: [1])
```
>>> isinstance(biggest_rel_change_major, (int, float))
True
```


Tests for q3_1 (points: [4])
```
>>> inventory.sort(0).column(0).item(0) == 25274
True
```


Tests for q3_2 (points: [0])
```
>>> all_different in {True, False}
True
```


Tests for q3_3 (points: [1, 1, 1])
```
>>> sales.num_rows == 8
True
```
```
>>> sales.num_columns == 4
True
```
```
>>> sales.labels == ('box ID', 'fruit name', 'count sold', 'price per fruit ($)')
True
```


Tests for q3_4 (points: [0])
```
>>> # We're asking for the number of *pieces* of fruit, not the
>>> # number of kinds of fruit or the number of boxes from which
>>> # there were sales.
>>> total_fruits_sold > 10
True
```
