```python
# Initialize Otter
import otter
grader = otter.Notebook("lab14.ipynb")
```

# Lab 14: Classification

Welcome to the last lab of the semester! 

This week's lab will focus on classification. 


**Helpful Resource:**

- [Python Reference](https://pages.mtu.edu/~lebrown/data1202-s24/reference/index.html): Cheat sheet of helpful array & table methods used in DATA 1202!

**Recommended Reading**: 

* [Classification](https://www.inferentialthinking.com/chapters/17/Classification.html)

**Submission**: Once you’re finished, run all cells besides the last one, select File > Save Notebook, and then execute the final cell. Then submit the downloaded zip file, that includes your notebook,  according to your instructor's directions.

First, set up the notebook by running the cell below.

```python
# Don't change this cell; just run it. 

import numpy as np
from datascience import *
import d8error

# These lines do some fancy plotting magic
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore')
from datetime import datetime
```

## 1. Bay Area School Coordinates with Classification

This lab is about k-Nearest Neighbors classification (k-NN). 

### Our Dearest Neighbors

Carol is trying classify students as either attendees of UC Berkeley or as attendees of Stanford University. To classify the students, Carol has access to the coordinates of the location they live during the school year. First, load in the `coordinates` table.

```python
# Just run this cell!
coordinates = Table.read_table('coordinates.csv')
coordinates.show(5)
```

As usual, let's investigate our data visually before performing any kind of numerical analysis.

```python
# Just run this cell!
coordinates.scatter("longitude", "latitude", group="school")
```

The locations of the points on this scatter plot look rather tightly grouped - run the following cell to see what they correspond to in terms of map locations.

```python
# Just run this cell!
colors = {"Berkeley":"blue", "Stanford":"red"}
t = Table().with_columns("lat", coordinates.column(0), 
                                      "lon", coordinates.column(1), 
                                      "color", coordinates.apply(colors.get, 2)
                        )
Circle.map_table(t, radius=5, fill_opacity=1)
```

**Question 1.1.** Let's begin implementing the k-Nearest Neighbors algorithm. Define the `distance` function, which takes in two arguments: an array of numerical features (`arr1`), and a different array of numerical features (`arr2`). The function should return the [Euclidean distance](https://en.wikipedia.org/wiki/Euclidean_distance) between the two arrays. Euclidean distance is often referred to as the straight-line distance formula that you may have learned previously. **(10 points)**



```python
def distance(arr1, arr2):
    ...

# Don't change/delete the code below in this cell
distance_example = distance(make_array(1, 2, 3), make_array(4, 5, 6))
distance_example
```

```python
grader.check("q1_1")
```

### Splitting the Dataset
We'll do two different kinds of things with the `coordinates` dataset:

1. We'll build a classifier using coordinates for which we know the associated label; this will teach it to recognize labels of similar coordinate values. This process is known as *training*.
2. We'll evaluate or *test* the accuracy of the classifier we build on data we haven't seen before.

As discussed in [Section 17.2](https://inferentialthinking.com/chapters/17/2/Training_and_Testing.html#training-and-testing), we want to use separate datasets for training and testing. As such, we split up our one dataset into two.

**Question 1.2.** Next, let's split our dataset into a training set and a test set. Since `coordinates` has 100 rows, let's create a training set with the first 75 rows and a test set with the remaining 25 rows. Remember that assignment to each group should be random, so we should shuffle the table first. **(10 points)**

*Hint:* As a first step we can shuffle all the rows, then use the `tbl.take` function to split up the rows for each table.


```python
np.random.seed(42) # set a random seed to make the shuffling repeatable
shuffled_table = ...
train = ...
test = ...

print("Training set:\t",   train.num_rows, "examples")
print("Test set:\t",       test.num_rows, "examples")
train.show(5), test.show(5);
```

```python
grader.check("q1_2")
```

**Question 1.3.** Assign `features` to an array of column name(s) (strings) of the features from the `coordinates` table.  Assign `classname` to an array of column name(s) (strings) from the `coordinates` table to be predicted. **(10 points)**

*Hint:* Which of the column names in the `coordinates` table are the features, and which of the column names correspond to the class we're trying to predict?

*Hint*: No need to modify any tables, just manually create an array of the feature names and class names!


```python
features = ...
print(features)

classname = ...
print(classname)
```

```python
grader.check("q1_3")
```

**Question 1.4.** Now define the `classify` function. This function should take in a `test_row` from a table like `test` and classify in using the k-Nearest Neighbors based on the correct `features` and the data in `train`. A refresher on k-Nearest Neighbors can be found [here](https://www.inferentialthinking.com/chapters/17/4/Implementing_the_Classifier.html). **(20+10 points)**


*Hint 1:* The `distance` function we defined earlier takes in arrays as input, so use the `row_to_array` function we defined for you to convert rows to arrays of features.

*Hint 2:* The skeleton code we provided iterates through each row in the training set.


```python
def row_to_array(row, features):
    """Converts a row to an array of its features."""
    arr = make_array()
    for feature in features:
        arr = np.append(arr, row.item(feature))
    return arr

def classify(test_row, k, train):
    test_row_features_array = row_to_array(test_row, features)
    distances = make_array()
    for train_row in train.rows:
        train_row_features_array = ...
        row_distance = ...
        distances = ...
    train_with_distances = ...
    nearest_neighbors = ...
    most_common_label = ...
    ...

# Don't modify/delete the code below
first_test = classify(test.row(0), 5, train)
first_test
```

```python
grader.check("q1_4")
```



**Question 1.5.** Define the function `three_classify` that takes a `row` from `test` as an argument and classifies the row based on using 3-Nearest Neighbors. Use this function to find the `accuracy` of a 3-NN classifier on the `test` set. `accuracy` should be a proportion (not a percentage) of the schools that were correctly predicted. **(10 points)**


*Hint:* You should be using a function you just created!

*Note:* Usually before using a classifier on a test set, we'd classify first on a "validation" set, which we then can modify our training set again if need be, before actually testing on the test set. You don’t need to do that for this question, but you will learn about this more in DATA 2201 and 3801.


```python
def three_classify(row):
    ...

test_with_prediction = ...
labels_correct = ...
accuracy = ...
accuracy
```

```python
grader.check("q1_5")
```

**Question 1.6.** There are 77 rows of Berkeley students and 23 rows of Stanford students in the `coordinates` table. If we used the entire `coordinates` table as the training set, what is the smallest value of k that would ensure that a k-Nearest Neighbor classifier would *always* predict Berkeley as the class? Assign the value to `k`. **(10 points)**


```python
k = ...
k
```

```python
grader.check("q1_6")
```

<!-- BEGIN QUESTION -->

**Question 1.7.** Why do we divide our data into a training and test set? What is the point of a test set, and why do we only want to use the test set once? Explain your answer in 3 sentences or less. **(10 points)**

*Hint:* Check out this [section](https://inferentialthinking.com/chapters/17/2/Training_and_Testing.html) in the textbook.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 1.8.** Why do we use an odd-numbered `k` in k-NN? Explain. **(10 points)**


_Type your answer here, replacing this text._

<!-- END QUESTION -->

## (OPTIONAL, NOT IN SCOPE): k-NN for Non-Binary Classification

**THIS IS NOT IN SCOPE**. There are no autograder tests for this or code for you to write. It just relies on the function `classify` in Question 1.4. Go ahead and read through this section and run the following cells!

In this class, we have taught you how to use the k-NN algorithm to classify data as one of two classes. However, much of the data you will encounter in the real world will not fall nicely into one of two categories. 

**How can we classify data with non-binary classes?** It turns out we can still use k-NN! That is, we find the distance between a point and all its neighbors, find the nearest neighbors, and take a majority vote among the neighbors to determine this point's class. 

The only difference is that now the neighboring points have more than two possible classes. This does introduce difficulty because now we have no way of guaranteeing that we will not encounter ties between classes. In the case that we do encounter a tie, we can just arbitrarily choose one of the classes.

In fact, you don't even have to modify the code you wrote before at all to enable multi-class classification!

Let's add some more data to our train table, this time for another class of students, students at San Jose Community College (SJCC).

```python
coordinates_multi = coordinates.with_rows([
                              [37.304346, -121.915401, "SJCC"],
                              [37.316275, -121.913879, "SJCC"],
                              [37.409435, -121.951379, "SJCC"],
                              [37.349387, -121.960771, "SJCC"],
                              [37.329083, -121.928479, "SJCC"],
                              [37.313017, -121.866730, "SJCC"],
                              [37.346525, -121.894767, "SJCC"],
                              [37.364157, -121.955717, "SJCC"],
                              [37.383362, -121.925776, "SJCC"],
                              [37.329545, -121.880639, "SJCC"]                             
])
```

```python
classify(coordinates_multi.row(0), 5, coordinates_multi)
```

```python
classify(coordinates_multi.row(91), 5, coordinates_multi)
```

```python
classify(coordinates_multi.row(105), 5, coordinates_multi)
```

Our classifier can classify rows as belonging to one of three classes!

Classification is one of the most important fields in statistics, data science, and machine learning. There are thousands of different classification algorithms and modifications of algorithms! There are many that you'll learn if you continue down the path of becoming a data scientist!

You're all done with Lab 14! :,)

**Important submission steps:** 
1. Run the tests and verify that they all pass.
2. Choose **Save Notebook** from the **File** menu, then **run the final cell**. 
3. Click the link to download the zip file.
4. Then submit the zip file to the corresponding assignment according to your instructor's directions. 

**It is your responsibility to make sure your work is saved before running the last cell.**

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: [0, 5, 5])
```
>>> type(distance_example) in set([float, np.float32, np.float64])
True
```
```
>>> np.isclose(distance_example, 5.196152422706632)
True
```
```
>>> np.isclose(distance(make_array(1, 1, 2, 3, 4, 5, 6, 8, 10), 
...                     make_array(1, 1, 3, 2, 5, 4, 8, 6, 10)),
...            3.4641016151)
True
```


Tests for q1_2 (points: [0, 0, 5, 1, 1, 1, 1, 1])
```
>>> # Double check that you have the correct number of rows for the `train` table.
>>> train.num_rows == 75
True
```
```
>>> # Double check that you have the correct number of rows for the `test` table.
>>> test.num_rows == 25
True
```
```
>>> train.num_rows + test.num_rows == coordinates.num_rows
True
```
```
>>> train.row(0).item(2) == 'Stanford'
True
```
```
>>> train.row(40).item(2) == 'Berkeley'
True
```
```
>>> np.isclose(train.row(73).item(0), 37.8648120000)
True
```
```
>>> np.isclose(test.row(10).item(1), -122.263019)
True
```
```
>>> np.isclose(test.row(22).item(0), 37.8638649999)
True
```


Tests for q1_3 (points: [0, 6, 4])
```
>>> type(features) == np.ndarray
True
```
```
>>> sorted(features) == ['latitude', 'longitude']
True
```
```
>>> all(classname == ['school'])
True
```


Tests for q1_4 (points: [0, 4, 4, 4, 4, 4])
```
>>> type(first_test) == str or type(first_test) == np.str_
True
```
```
>>> sorted_coordinates = coordinates.sort("school")
>>> classify(sorted_coordinates.row(15), 3, sorted_coordinates.take(np.arange(50, 100))) == 'Berkeley'
True
```
```
>>> sorted_coordinates = coordinates.sort("school")
>>> classify(sorted_coordinates.row(25), 3, sorted_coordinates.take(np.arange(50, 100))) == 'Berkeley'
True
```
```
>>> sorted_coordinates = coordinates.sort("school")
>>> classify(sorted_coordinates.row(98), 3, sorted_coordinates.take(np.arange(45, 98))) == 'Stanford'
True
```
```
>>> sorted_coordinates = coordinates.sort("school")
>>> classify(sorted_coordinates.row(0), 5, sorted_coordinates.take(np.arange(25,100))) == 'Berkeley'
True
```
```
>>> sorted_coordinates = coordinates.sort("school")
>>> classify(sorted_coordinates.row(99), 5, sorted_coordinates.take(np.arange(25,95))) == 'Berkeley'
True
```


Tests for q1_5 (points: [4, 6])
```
>>> sorted_coordinates = coordinates.sort("school")
>>> classify(sorted_coordinates.row(29), 3, train) == three_classify(sorted_coordinates.row(29))
True
```
```
>>> # You should expect to see really high accuracy
>>> 0.90 <= accuracy <= 1
True
```


Tests for q1_6 (points: [0, 10])
```
>>> # `k` should be an int
>>> type(k) == int
True
```
```
>>> import hashlib 
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(k) == '67c6a1e7ce56d3d6fa748ab6d9af3fd7'
True
```
