```python
# Initialize Otter
import otter
grader = otter.Notebook("lab13.ipynb")
```

# Lab 13: Variability and the Normal Distribution

In this lab you will practice calculating variance and standard deviation, and converting values to standard units. 

You will use these skills to compare grades in a course and you will use Chebyshev’s inequality to predict how hard the students should work in order to rank in the top 5% of the class. Finally, you will use confidence intervals to help college administrators plan for next quarter by predicting the enrollment in a new course.

The topics from this lab are covered in the following readings:

* [CIT 14.2](https://www.inferentialthinking.com/chapters/14/2/Variability.html): Variability
* [CIT 14.3](https://www.inferentialthinking.com/chapters/14/3/SD_and_the_Normal_Curve.html): The Standard Deviation (SD) and the Normal Curve 
* [CIT 14.4](https://www.inferentialthinking.com/chapters/14/4/Central_Limit_Theorem.html): The Central Limit Theorem
* [CIT 14.5](https://www.inferentialthinking.com/chapters/14/5/Variability_of_the_Sample_Mean.html): The Variability of the Sample Mean
* [CIT 14.6](https://inferentialthinking.com/chapters/14/6/Choosing_a_Sample_Size.html): Choosing a Sample Size

As usual, run the cell below to prepare the lab and the automatic tests.

```python
# Run this cell to set up the notebook, but please don't change it.

# These lines import the Numpy and Datascience modules.
import numpy as np
from datascience import *

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', (FutureWarning, np.VisibleDeprecationWarning))

import d8error
```

## 1. Comparing Grades Using Standard Units

Two of your friends, Jasmine and Zelong, just took their midterms. Jasmine took her CS 1122 midterm and Zelong took his MATH 1161 midterm. Jasmine received a **B+** on her midterm (87%) and Zelong received an **A-** (92%). Jasmine claims that while she received a lower grade on her midterm, she actually did better (relative to the rest of the class) than Zelong. Zelong disagrees. Knowing that you are taking DATA 1202, your two friends come to you to settle their argument. 

They show you two DataFrames: `cs_midterm` and `math_midterm` that represent the grades for their classes. Both exams are out of 100 points.  Each DataFrame has a column called `'Student'` with student ID numbers and `'Score'` with the midterm scores.

```python
# Jasmine's exam 
cs_midterm = Table.read_table("cs_scores.csv")
cs_midterm.hist("Score", bins=range(0, 101, 4))
plt.title('CS 1122 Midterm Scores')
Jasmine_score = 87 
print("Jasmine Score (CS 1122): " + str(Jasmine_score))


# Zelong's exam 
math_midterm = Table.read_table("math_scores.csv")
math_midterm.hist("Score", bins=range(0, 101, 4))
plt.title('MA 1161 Midterm Scores')
Zelong_score = 92
print("Zelong Score (MA 1161): " + str(Zelong_score))
```

You know that instead of comparing their actual scores, you should first convert their scores into **standard units**. Recall from Lecture 21 that if $x$ is a numerical variable, and $x_i$ is one value of the variable, then

$$x_{i } \text{ in standard units } = \frac{x_i - \text{mean of $x$}}{\text{SD of $x$}}$$

represents $x_i$ in **standard units** – the number of standard deviations $x_i$ is above the mean.

To compute the midterm score in standard units for each friend, we need to:
1. Compute the **average** grade for each class. 
2. Compute the **standard deviation** (SD) of the midterm scores for each class. We *could* use `np.std`, but we will write our own function to do that.

Note that standard deviation is the square root of variance. So, we'll proceed by defining a function that computes the **variance** first. Recall that the variance is the mean squared deviation from the average:

$$\text{variance} = \frac{(\text{value}_1 - \text{average})^2 + (\text{value}_2 - \text{average})^2 +...+ (\text{value}_n - \text{average})^2}{n}$$
where $n$ is the number of values (e.g. number of exam scores, in our case).

**Question 1.1.** Complete the function `compute_variance`. It takes as input an array (`data`) and returns the variance as a single number. 

Then, use the `compute_variance` function to compute the variance of the two classes' midterm scores, and assign them to the two specified variable names.

**Do not** use `np.std` in your solution. Instead, use the above formula for variance as guidance.

```python
def compute_variance(data):
    n = ...
    average = ...
    sum_squared_diff = ...
    variance = sum_squared_diff / n 
    return variance


cs_midterm_var = ...
print("Variance of CS 1122 midterm: " + str(cs_midterm_var))

math_midterm_var = ...
print("Variance of MA 1161 midterm: " + str(math_midterm_var))
```

```python
grader.check("q1_1")
```

**Question 1.2.** Now that we have a function that computes the variance, we want to write a function that computes the standard deviation. Fill in the missing code to complete the function `compute_sd`. It takes as input an *array of numbers* (`data`) and returns the *standard deviation* as a single number. 

Then, use the `compute_sd` function to compute the standard deviation of scores of the two midterms.

_Hint:_ Your implementation of `compute_sd` should only take ***one line***, that involves both the `return` keyword and the function `compute_variance`.

```python
def compute_sd(data):
    ...

cs_midterm_sd = ...
print("Standard Deviation of CS 1122 midterm: " + str(cs_midterm_sd))

math_midterm_sd = ...
print("Standard Deviation of MA 1161 midterm: " + str(math_midterm_sd))
```

```python
grader.check("q1_2")
```

**Question 1.3.** Now that you can compute the standard deviation, you are equipped to write a function that converts a given score to standard units. Fill in the missing code to complete the function `compute_su`. It takes in a *score* (`score`), the *average score* (`avg`), and the *standard deviation* (`sd`), and returns the score in standard units.

Then, use the `compute_su` function to transform the scores earned by each friend into standard units.

**Warning**: Be careful with the order of operations!

```python
def compute_su(score, avg, sd):
    standard_units = ...
    return standard_units

Jasmine_su = ...
print("Jasmine's Score in Standard Units: " + str(Jasmine_su))

Zelong_su = ...
print("Zelong's Score in Standard Units: " + str(Zelong_su))
```

```python
grader.check("q1_3")
```

**Question 1.4.** Jasmine's score *is* higher than Zelong's score when we convert to standard units, which can be seen as evidence that she did better on her exam relative to her classmates than Zelong did relative to his. 

Another way to measure their relative performances is to calculate, for both Jasmine and Zelong individually, the proportion of students they scored higher than (or the same as). Comparing Jasmine's proportion to Zelong's proportion will give us another way of measuring who did better relative to their classmates. Calculate Jasmine's proportion and Zelong's proportion below. (This will require looking at both `cs_midterm` and `math_midterm`.)

```python
Jasmine_proportion = ...
print("Jasmine's Proportion: " + str(Jasmine_proportion))

Zelong_proportion = ...
print("Zelong's Proportion: " + str(Zelong_proportion))
```

```python
grader.check("q1_4")
```

## 2. Chebyshev's Inequality and Normal Curves

Lets look at the histograms of the scores of the two midterms again.

```python
cs_midterm.hist("Score", bins=range(0, 101, 4))
plt.title('CS 1122 Midterm Scores');

math_midterm.hist("Score", bins=range(0, 101, 4))
plt.title('MA 1161 Midterm Scores');
```

**Question 2.1.** Which of the two graphs roughly resembles a normal curve? Assign the variable `q2_1` to either 1, 2, 3, or 4.

1. Only the first graph (distribution of CS 1122 midterm scores) is roughly normal.
2. Only the lower graph (distribution of MA 1161 midterm scores) is roughly normal.
3. Both graphs are roughly normal.
4. Neither graph is roughly normal.

Remember all normal curves have the following characteristics:

* The mean (average) is always in the center of a normal curve.
* A normal curve has only one mode (peak).

```python
q2_1 = ...
```

```python
grader.check("q2_1")
```

**Question 2.2.**
By looking at the distribution of MA 1161 midterm scores above, rank the following values in order **from smallest to largest**.

1. The mean score.
2. The median score.
3. The most common score (the mode).

Set variable `q2_2` to a list containing the numbers 1, 2, 3 **in the appropriate order**. Don't compute any of these values manually!

```python
q2_2 = ...
```

```python
grader.check("q2_2")
```

### Recap: Chebyshev's Inequality 

Chebyshev's Inequality states that *no matter what the shape of the distribution*, the proportion of values in the range "mean +/- z SDs" is

at least $$1 - \frac{1}{z^2}$$

**It's important to note that these are lower bounds, not approximations:** 75% of the data is guaranteed to lie within plus or minus of 2 standard deviations of the mean, but 100% of the data might also lie within plus or minus 2 standard deviations of the mean. 

### On the other hand...

**If we know that our data forms a normal curve**, the standard deviation is even more informative.

|Percent in Range | All Distributions (via Chebyshev's Inequality) | Normal Distributions|
|---|---|---|
|$\text{average} \pm 1 \ \text{SD}$ | $\geq 0\%$ | $\approx 68\%$ |
|$\text{average} \pm 2\text{SDs}$ | $\geq 75\%$ | $\approx 95\%$ |
|$\text{average} \pm 3\text{SDs}$ | $\geq 88\%$ | $\approx 99.73\%$ |

**Note that for a normal distribution, the numbers in the last column of the DataFrame above are approximations, not lower bounds.**  
* If the distribution is perfectly normal, then 68% of the data (not more, not less) will lie between plus and minus one standard deviation of the mean. 
* Additionally because a normal curve is symmetric, we know that 34% of the data lies between the average and the average plus one standard deviation. 


**Question 2.3.**  Jasmine, who is Pre-Med, really wanted to score in the top 5% of the class. But before taking the exam, she did not know if the scores would be normally distributed or not. 

Without making any assumptions about the distribution of scores, how many standard deviations above the mean would she have needed to score to **guarantee** that she fell in the top 5% of the the class? Set the variable `q2_3` to either 1, 2, 3, or 4, depending on your answer. 


1. Jasmine would need to score roughly 4.5 standard deviations above the average to guarantee being in the top five percent. Using Chebyshev's inequality, setting $z = \sqrt{20} \approx 4.5$  gives that 95% of the data will lie between plus or minus 4.5 SDs. If Jasmine scores above 4.5 SDs, then she is guaranteed to have scored better than 95% of the other students. 

2. Jasmine would need to score above 2 SDs. Since 95% of the data falls between plus or minus 2 SDs, if Jasmine scores above 2 SDs, she is guaranteed to score above 95% of the class. 

3. Jasmine would need to score slightly less than 2 SDs. 50% of the class will have scored below the average. Which means that if Jasmine scores 2 standard deviations above the average she'll have scored higher than 50% + (95% / 2) = 97.5%. 

4. No matter how many standard deviations above the mean Jasmine scores, there is no guarantee that she will score in the top 5% of the class. 

```python
q2_3 = ...
```

```python
grader.check("q2_3")
```

**Question 2.4.** Now, assuming that the scores for the exam will be normally distributed (as many exams are), what is the minimum number of standard deviations above the mean Jasmine would have needed to score to **guarantee** that she fell in the top 5% of the class? Set variable `q2_4` to either 1, 2, 3, or 4, depending on your answer. 


1. Jasmine would need to score roughly 4.5 standard deviations above the average to guarantee being in the top five percent. Using Chebyshev's inequality, setting $z = \sqrt{20} \approx 4.5$  gives that 95% of the data will lie between plus or minus 4.5 SDs. If Jasmine scores above 4.5 SDs, then she is guaranteed to have scored better than 95% of the other students. 

2. Jasmine would need to score above 2 SDs. Since 95% of the data falls between plus or minus 2 SDs, if Jasmine scores above 2 SDs, she is guaranteed to score above 95% of the class. 

3. Jasmine would need to score slightly less than 2 SDs. 50% of the class will have scored below the average. Which means that if Jasmine scores 2 standard deviations above the average she'll have scored higher than 50% + (95% / 2) = 97.5% of the class.

4. No matter how many standard deviations above the mean Jasmine scores, there is no guarantee that she will score in the top 5% of the class.

```python
q2_4 = ...
```

```python
grader.check("q2_4")
```

## 3. Choosing Sample Size

A new class is being offered at MTU and the administration wants to know how many students will be taking the class so they know how big of a classroom / number of sections it will need. To take the class, a student must have satisfied the prerequisites first. 

The administration knows there are 900 students who have completed the prerequisites and are eligible to take the class, but they don't know how many will actually enroll. They don't have the resources to ask each eligible student whether they are going to enroll, so they decide to ask a sample of the students, but they don't know how many students to ask. They want the width of their confidence interval to be **at most 10 students**. 

For example, if the results of their sample concluded that with 95% confidence between 200 and 210 students would take the class, the administration would be happy with that sample. However if the results of the sample concluded that with 95% confidence between 200 and 300 students would take the class, the sample would not have been informative enough because that range is too wide. We are going to help determine how big of a sample the administration should take. 

The population parameter we are interested in measuring is the proportion of eligible students who will take the class. We will estimate this using a sample statistic, the proportion of eligible students in the sample who plan to take the class. 

So **where do we start**?

The Central Limit Theorem tells us that regardless of the distribution of our population, the distribution of the sample mean or sum will always be roughly normal. Fortunately, our sample statistic (the proportion of eligible students who will take the class) is also a sample mean, because proportions are just means of 0s and 1s. Let's run a simulation to see this for ourselves. 

Below is the data for the whole population of eligible students. (If the administration had the resources to ask every student whether they were going to take the class, this is what they would see. `0` means they won't take the class and `1` means they will.)

```python
population = Table.read_table("population.csv")
population
```

```python
population.hist('Planning on taking')
```

**Question 3.1.** Below is partially implemented code to run a simulation. The simulation will repeatedly take samples of size `sample_size` (without replacement) from `population` and calculate the proportion of students who plan on taking the class. Fill in the missing parts. 

```python
def simulation(population, num_iterations, sample_size): 
    results = make_array()
    for i in np.arange(num_iterations):
        sampled = ...
        proportion_taking_class = ...
        results = ...

    Table().with_column('Proportion', results).hist('Proportion')
```

```python
simulation(population, 10000, 40)
```

```python
grader.check("q3_1")
```

**Question 3.2.** Does the distribution of the sample proportion look more like a normal curve or more like the population distribution? 

1. More like a normal curve.

2. More like the original population.

```python
q3_2 = ...
```

```python
grader.check("q3_2")
```

We also know that as we increase the sample size, the standard deviation of our sample proportion's distribution will decrease. Again we decide to run a simulation to double check. Run the following cell to see how the distribution of the sample proportion changes as we increase the size of our sample. **It might take a while to run.** 

```python
simulation(population, 10000, 40)
simulation(population, 10000, 120)
simulation(population, 10000, 360)
```

This trend is expressed by the formula
$$ \text{SD of Distribution of Possible Sample Means} = \frac{\text{Population SD}}{\sqrt{\text{sample size}}}$$

Since proportions are means, we can use this formula to find the sample size we need to get a desired standard deviation of the sample proportion, and thus a certain confidence interval for that sample proportion. However, before taking our sample, we don't have any way of knowing the standard deviation of our population. Lecture 23 and  the [CIT](https://inferentialthinking.com/chapters/14/6/Choosing_a_Sample_Size.html) textbook include ways to get around this problem; here we will use the actual population standard deviation. 

**Question 3.3.** The administration wants the confidence interval to have a width of **10 students**, but we have been calculating the proportion of eligible students who are planning on taking the class. Fortunately, there is a relationship between a count and a proportion that will allow us to convert between these.

Using the total number of students who are eligible to take the class, determine what proportion of that number equals 10 students, and save the result as `width_as_proportion`. For example, if 500 students are eligible to take the class, then 10 students as a proportion is 0.02.

```python
num_eligible_students = ...
print('Number of eligible students:', num_eligible_students)
width_as_proportion = ...
print('Desired confidence interval width, as a proportion:', width_as_proportion)
```

```python
grader.check("q3_3")
```

**Question 3.4.** Now let's calculate the standard deviation that our distribution of sample proportions would need to have in order for our 95% confidence interval to have a width of `width_as_proportion`. Set the variable `target_sd` to this value. 

Remember that for a normal distribution, 95% of the data lies between *plus and minus* 2 SDs of the mean.

```python
target_sd = ...
target_sd
```

```python
grader.check("q3_4")
```

**Question 3.5.** We also need to calculate the standard deviation of the total population. Calculate this value using the`compute_sd` function that you wrote earlier and store it in the variable `population_sd`.

```python
population_sd = ...
population_sd
```

```python
grader.check("q3_5")
```

**Question 3.6.** Now calculate the required **sample size** and store your result as `req_sample_size`. Recall that

$$ \text{SD of Distribution of Possible Sample Means} = \frac{\text{Population SD}}{\sqrt{\text{sample size}}}$$

You have already calculated $\text{SD of Distribution of Possible Sample Means}$, in Question 3.4, and $\text{Population SD}$, in Question 3.5.

```python
req_sample_size = ...
req_sample_size
```

```python
grader.check("q3_6")
```

**Question 3.7.** Our required sample size is bigger than our entire population. For each part, say whether it is `True` or `False`.

1. The administration will have to settle for a wider interval to get 95% confidence.
2. Sampling with replacement will be a feasible way to determine the information the administration needs.
3. The administration will have to settle for a lower degree of confidence to get an interval of width 10.
4. We should increase the size of the population until the sample size is smaller than the size of the population.

Set each variable below to either `True` or `False`.

```python
statement_1 = ...
statement_2 = ...
statement_3 = ...
statement_4 = ...
```

```python
grader.check("q3_7")
```

## 4. Submission

**Important submission steps:** 
1. Run the tests and verify that they all pass.
2. Choose **Save Notebook** from the **File** menu, then **run the final cell**. 
3. Click the link to download the zip file.
4. Then submit the zip file to the corresponding assignment according to your instructor's directions. 

**It is your responsibility to make sure your work is saved before running the last cell.**

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: [2, 2, 2])
```
>>> np.isclose(cs_midterm_var, 173.20529599999983)
True
```
```
>>> np.isclose(math_midterm_var, 84.26559999999982)
True
```
```
>>> np.isclose(compute_variance(make_array(1, 2, 3, 4, 5)), 2.0)
True
```


Tests for q1_2 (points: [2, 2, 2])
```
>>> np.isclose(cs_midterm_sd, 13.160748306992268)
True
```
```
>>> np.isclose(math_midterm_sd, 9.179629622157956)
True
```
```
>>> np.isclose(compute_sd(make_array(1, 2, 3, 4, 5)), 1.4142135623730951)
True
```


Tests for q1_3 (points: [3, 3])
```
>>> np.isclose(Jasmine_su, 1.6676863266459616)
True
```
```
>>> np.isclose(Zelong_su, 0.38345773684630546)
True
```


Tests for q1_4 (points: [3, 3])
```
>>> Jasmine_proportion == 0.064
True
```
```
>>> Zelong_proportion == 0.472
True
```


Tests for q2_1 (points: [1.5, 1.5, 1.5, 1.5])
```
>>> 0 < q2_1 < 5
True
```
```
>>> q2_1 != 2
True
```
```
>>> q2_1 != 3
True
```
```
>>> q2_1 != 4
True
```


Tests for q2_2 (points: [1, 1, 1, 1, 1, 1])
```
>>> type(q2_2) == np.ndarray
True
```
```
>>> any(q2_2 != [2, 1, 3])
True
```
```
>>> any(q2_2 != [2, 3, 1])
True
```
```
>>> any(q2_2 != [1, 3, 2])
True
```
```
>>> any(q2_2 != [3, 1, 2])
True
```
```
>>> any(q2_2 != [3, 2, 1])
True
```


Tests for q2_3 (points: [1.5, 1.5, 1.5, 1.5])
```
>>> 0 < q2_3 < 5
True
```
```
>>> q2_3 != 2
True
```
```
>>> q2_3 != 3
True
```
```
>>> q2_3 != 4
True
```


Tests for q2_4 (points: [2, 2, 2, 2])
```
>>> 0 < q2_4 < 5
True
```
```
>>> q2_4 != 2
True
```
```
>>> q2_4 != 1
True
```
```
>>> q2_4 != 4
True
```


Tests for q3_1 (points: [7])
```
>>> callable(simulation)
True
```


Tests for q3_2 (points: [4, 3])
```
>>> 0 < q3_2 < 3 
True
```
```
>>> q3_2 != 2
True
```


Tests for q3_3 (points: [7])
```
>>> np.isclose(width_as_proportion, 0.011111111111111112)
True
```


Tests for q3_4 (points: [7])
```
>>> np.isclose(target_sd, 0.002777777777777778)
True
```


Tests for q3_5 (points: [7])
```
>>> np.isclose(population_sd, 0.4582575694955866)
True
```


Tests for q3_6 (points: [7])
```
>>> np.isclose(req_sample_size, 27216)
True
```


Tests for q3_7 (points: [2, 2, 2, 2])
```
>>> statement_1 == True
True
```
```
>>> statement_2 == False
True
```
```
>>> statement_3 == True
True
```
```
>>> statement_4 == False
True
```
