```python
# Initialize Otter
import otter
grader = otter.Notebook("hw12.ipynb")
```

# HW 12: Normal Distribution and Variability of Sample Means

**Helpful Resource:**

- [Python Reference](https://pages.mtu.edu/~lebrown/data1202-s24/reference/index.html): Cheat sheet of helpful array & table methods used in DATA 1202!

**Recommended Reading**: 

* [Normal distribution](https://www.inferentialthinking.com/chapters/14/3/SD_and_the_Normal_Curve.html)
* [Central Limit Theorem](https://inferentialthinking.com/chapters/14/4/Central_Limit_Theorem.html)


Please complete this notebook by filling in the cells provided. Before you begin, execute the following cell to setup the notebook by importing some helpful libraries. Each time you start your server, you will need to execute this cell again.

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. **Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook!** For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!


**Note: This homework has hidden tests on it. That means even though the tests may say 100% passed, it doesn't mean your final grade will be 100%. We will be running more tests for correctness once everyone turns in the homework.**


Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. 

You should start early so that you have time to get help if you're stuck.

```python
# Run this cell, but please don't change it.

# These lines import the Numpy and Datascience modules.
import numpy as np
from datascience import *

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore')

import d8error
```

# 1. Normal Distributions

When we visualize the distribution of a sample, we are often interested in the mean and the standard deviation of the sample (for the rest of this lab, we will abbreviate “standard deviation” as “SD”). These two sample statistics can give us a bird’s eye view of the distribution -- by letting us know where the distribution sits on the number line and how spread out it is, respectively. 

We want to check if the data is linearly related, so we should look at the data.

<!-- BEGIN QUESTION -->

**Question 1.1.** The next cell loads the table `births` from lecture, which is a large random sample of US births and includes information about mother-child pairs. 

Plot the distribution of mother’s ages from the table. Don’t change the last line, which will plot the mean of the sample on the distribution itself.


```python
births = Table.read_table('baby.csv')
...

# Do not change anything below this line
plots.scatter(np.mean(births.column("Maternal Age")), 0, color='red', zorder=2, s=50);
plots.ylim(-0.005, 0.07);
```

<!-- END QUESTION -->

From the plot above, we can see that the mean is the center of gravity or balance point of the distribution. If you cut the distribution out of cardboard, and then placed your finger at the mean, the distribution would perfectly balance on your finger. Since the distribution above is right skewed (which means it has a long right tail), we know that the mean of the distribution is larger than the median, which is the “halfway” point of the data. Conversely, if the distribution had been left skewed, we know the mean would be smaller than the median.

<br>Run the following cell to compare the mean (red) and median (blue) of the distribution of mothers ages.


```python
births.hist("Maternal Age")
plots.scatter(np.mean(births.column("Maternal Age")), 0, color='red', zorder=2, s=50, label="Mean")
plots.scatter(np.median(births.column("Maternal Age")), 0, color='blue', zorder=2, s=50, label="Median")
plots.legend()
plots.ylim(-0.005, 0.07);
```

We are also interested in the standard deviation of mother’s ages. The SD gives us a sense of how variable mothers' ages are around the average mothers' age. If the SD is large, then the mothers' ages should spread over a large range from the mean. If the SD is small, then the mothers' ages should be tightly clustered around the average mother age. 

**The SD of an array is defined as the root mean square of the deviations where the deviations are the difference between the original elements and their average.**

You can read more about Variability and SD in [Chapter 14.2](https://inferentialthinking.com/chapters/14/2/Variability.html). Fun fact! $\sigma$ (Greek letter sigma) is used to represent the SD and $\mu$ (Greek letter mu) is used for the mean. 

**Question 1.2.** Assign `age_mean` to the mean of the maternal ages and `age_sd` to the SD of the maternal ages. Then run the cell below to see the width of one SD (blue) from the sample mean (red) plotted on the histogram of maternal ages.


```python
age_mean = ...
age_sd = ...
births.hist("Maternal Age")

plots.scatter(age_mean, 0, color='red', zorder=2, s=50);
plots.scatter(age_mean+age_sd, 0, marker='^', color='blue', zorder=2, s=50);
plots.scatter(age_mean-age_sd, 0, marker='^', color='blue', zorder=2, s=50);
plots.ylim(-0.005, 0.07);
```

```python
grader.check("q1_2")
```

In this histogram, the standard deviation is not easy to identify just by looking at the graph.

However, the distributions of some variables allow us to easily spot the standard deviation on the plot. For example, if a sample follows a *normal distribution*, the standard deviation is easily spotted at the point of inflection (the point where the curve begins to change the direction of its curvature) of the distribution. 

**Question 1.3.** Fill in the following code to examine the distribution of maternal heights, which is roughly normally distributed. We’ll plot the standard deviation on the histogram, as before - notice where one standard deviation (blue) away from the mean (red) falls on the plot.


```python
height_mean = ...
height_sd = ...
births.hist("Maternal Height", bins=np.arange(55,75,1))

plots.scatter((height_mean), 0, color='red', zorder=2, s=50);
plots.scatter(height_mean+height_sd, 0, marker='^', color='blue', zorder=2, s=50);
plots.scatter(height_mean-height_sd, 0, marker='^', color='blue', zorder=2, s=50);
plots.ylim(-0.01, 0.16);
```

```python
grader.check("q1_3")
```

We don’t always know how a variable will be distributed, and making assumptions about whether or not a variable will follow a normal distribution is dangerous. However, the **Central Limit Theorem** defines one distribution that always follows a normal distribution. **The distribution of the *sums* and *means* of all large random samples drawn with replacement from a single distribution (regardless of the distribution’s original shape) will be normally distributed.** Remember that the Central Limit Theorem refers to the distribution of a *statistic* calculated from a distribution, not the distribution of the original sample or population. 

The next section will explore distributions of sample means, and you will see how the standard deviation of these distributions depends on sample sizes.

# 2. Variability of the Sample Mean

By the [Central Limit Theorem](https://www.inferentialthinking.com/chapters/14/4/Central_Limit_Theorem.html), the probability distribution of the mean or sum of a large random sample is roughly normal. The bell curve is centered at the population mean. Some of the sample means are higher and some are lower, but the deviations from the population mean are roughly symmetric on either side, as we have seen repeatedly. 

In our simulations, we also noticed that the means of larger samples tend to be more tightly clustered around the population mean than means of smaller samples. In this section, we will quantify the [variability of the sample mean](https://www.inferentialthinking.com/chapters/14/5/Variability_of_the_Sample_Mean.html) and develop a relation between the variability and the sample size.

Let's take a look at the salaries of employees of the City of San Francisco in 2014. The mean salary reported by the city government was about $75,463.92.

*Note: If you get stuck on any part of this hw, please refer to [Chapter 14 of the textbook](https://www.inferentialthinking.com/chapters/14/Why_the_Mean_Matters.html).*

```python
salaries = Table.read_table('sf_salaries_2014.csv').select("salary")
salaries.show(5)
```

```python
salary_mean = np.mean(salaries.column('salary'))
print('Mean salary of San Francisco city employees in 2014: ', salary_mean.round(2))
```

```python
salaries.hist('salary', bins=np.arange(0, 250000+10000*2, 10000))
plots.scatter(salary_mean, 0, color='red', zorder=2, s=50)
plots.title('2014 salaries of city of SF employees')
plots.ylim(-0.0000005, 0.000014);
```

Clearly, the population *does not* follow a normal distribution. Keep that in mind as we progress through these exercises.

Let's take random samples **with replacement** and look at the probability distribution of the sample mean. As usual, we will use simulation to get an empirical approximation to this distribution.

**Question 2.1.** Define a function `one_sample_mean`. Its arguments should be `table` (the name of a table), `label` (the label of the column containing the variable), and `sample_size`(the number of employees in the sample). It should sample with replacement from the table and
return the mean of the `label` column of the sample.


```python
def one_sample_mean(table, label, sample_size):
    new_sample = ...
    new_sample_mean = ...
    ...
```

```python
grader.check("q2_1")
```

**Question 2.2.** Use `one_sample_mean` to define a function `simulate_sample_mean`. The arguments are the name of the `table`, the `label` of the column containing the variable, the `sample_size`, and the number of simulations (`repetitions`).

The function should sample with replacement from the table and calculate the mean of each sample (*Hint, use your function from Q2.1*). It should save the sample means in an array called `means`. The remaining code in the function displays an empirical histogram of the sample means.

*Note: We round all values to 2 decimal places because we are working with salaries.*


```python
"""Empirical distribution of random sample means"""

def simulate_sample_mean(table, label, sample_size, repetitions, verbose=1):
    
    means = make_array()

    for i in np.arange(repetitions):
        new_sample_mean = ...
        means = ...

    sample_means = Table().with_column('Sample Means', means)
    
    # Display empirical histogram and print all relevant quantities – don't change this!
    sample_means.hist(bins=20)
    plots.xlabel('Sample Means')
    plots.title('Sample Size {sample_size}; {repetitions} Resamples'.format(sample_size=sample_size, repetitions=repetitions))
    if verbose == 1:
        print("Sample size: ", sample_size)
        print("Population mean:", np.mean(table.column(label)).round(2))
        print("Average of sample means: ", np.mean(means).round(2))
        print("Population SD:", np.std(table.column(label)).round(2))
        print("SD of sample means:", np.std(means).round(2))
    return np.std(means)
```

```python
grader.check("q2_2")
```

Think very carefully that you've implemented the function above correctly. If you haven't implemented it correctly, the rest of the hw won't work properly, so this step is crucial.

 In the following cell, we will create a sample of size 100 from `salaries` and graph it using our new `simulate_sample_mean` function.
 
 *Hint: You should see a distribution similar to something we've been talking about. If not, check your function.*

```python
simulate_sample_mean(salaries, 'salary', 100, 10000, 1) 
plots.xlim(50000, 100000);
```

Notice that our distribution of *sample means* looks approximately normal! 

Did it matter if our original distribution was normally distributed?

<!-- BEGIN QUESTION -->

**Question 2.3.** Simulate two sample means, one for a sample of 400 salaries and one for a sample of 625 salaries. In each case, perform 10,000 repetitions. Don't worry about the `plots.xlim` or `plots.ylim` lines – they just make sure that all of the plots have the same x-axis and y-axis, respectively. 


```python
simulate_sample_mean(..., ..., ..., ...)
plots.xlim(50000, 100000)
plots.ylim(0, 0.0002)
plots.show();
print('\n')
simulate_sample_mean(..., ..., ..., ...)
plots.xlim(50000, 100000)
plots.ylim(0, 0.0002)
plots.show();
```

<!-- END QUESTION -->

**Question 2.4.** Assign `q2_4` to an array of numbers corresponding to true statement(s) about the plots from 2.3. 

*Hint*: If a distribution is bell-shaped, what type of curve does it follow? 

1. We see the Central Limit Theorem (CLT) in action because the distributions of the sample means are bell-shaped.
2. We see the Law of Averages in action because the distributions of the sample means look like the distribution of the population.
3. One of the conditions for CLT is that we have to draw a small random sample with replacement from the population.
4. One of the conditions for CLT is that we have to draw a large random sample with replacement from the population.
5. One of the conditions for CLT is that the population must be normally distributed.


```python
q2_4 = ...
```

```python
grader.check("q2_4")
```

**Question 2.5.** Assign `q2_5` to an array of numbers corresponding to true statement(s) about the plots from 2.3.

1. Both plots in 2.3 are roughly centered around the population mean.
2. Both plots in 2.3 are roughly centered around the mean of a particular sample.
3. The distribution of sample means for sample size 625 has less variability than the distribution of sample means for sample size 400.
4. The distribution of sample means for sample size 625 has more variability than the distribution of sample means for sample size 400.


```python
q2_5 = ...
```

```python
grader.check("q2_5")
```

Below, we'll look at what happens when we take an **increasing number of resamples of a fixed sample size.** Notice what number in the code changes, and what stays the same. How does the distribution of the resampled means change?

```python
simulate_sample_mean(salaries, 'salary', 100, 500, 1)
plots.xlim(50000, 100000)
plots.ylim(0, 0.00009);
```

```python
simulate_sample_mean(salaries, 'salary', 100, 1000, 1)
plots.xlim(50000, 100000);
plots.ylim(0, 0.00009);
```

```python
simulate_sample_mean(salaries, 'salary', 100, 5000, 1)
plots.xlim(50000, 100000);
plots.ylim(0, 0.00009);
```

```python
simulate_sample_mean(salaries, 'salary', 100, 10000, 1)
plots.xlim(50000, 100000);
plots.ylim(0, 0.00009);
```

What did you notice about the distributions of sample means in the four histograms above? 

**Question 2.6.** Assign the variable `SD_of_sample_means` to the integer corresponding to your answer to the following question:

When I increase the number of *resamples* that I take, for a *fixed* sample size, the SD of my sample means will...

1. Increase
2. Decrease
3. Stay about the same
4. Vary widely


```python
SD_of_sample_means = ...
```

```python
grader.check("q2_6")
```

**Question 2.7.** Let's think about how the relationships between population SD, sample SD, and SD of sample means change with varying sample size. Which of the following is true? Assign the variable `pop_vs_sample` to an array of integer(s) that correspond to true statement(s).

*Hint 1:* The sample SD is different from the SD of sample means. 

1. Sample SD gets smaller with increasing sample size.
2. Sample SD gets larger with increasing sample size.
3. Sample SD becomes more consistent with population SD with increasing sample size.
4. SD of sample means gets smaller with increasing sample size.
5. SD of sample means gets larger with increasing sample size.
6. SD of sample means stays the same with increasing sample size.


```python
pop_vs_sample = ...
```

```python
grader.check("q2_7")
```

## 3. Key Concepts 

Let's review some key concepts from this past week. 


**Question 3.1** How do we convert the value 34 to standard units if it comes from a dataset where the mean is 20 and the standard deviation is 3? Assign `q3_1` to either 1, 2, 3, or 4.

1. $\displaystyle \frac{20 - 34}{3}$
2. $\displaystyle \frac{34 - 20}{3}$
3. $\displaystyle \frac{34 - 20}{\sqrt{3}}$
4. $\displaystyle \frac{(20-34)^2}{3}$  


```python
q3_1 = ...
```

```python
grader.check("q3_1")
```

**Question 3.2** According to Chebyshev's inequality, for any dataset, at least half of the data falls within how many standard deviations of the mean? Assign the **smallest** correct answer to `q3_2`.

1. 1.35
2. 1.40
3. 1.45
4. 1.50


```python
q3_2 = 3  # SOLUTION
```

```python
grader.check("q3_2")
```

**Question 3.3** Assign `q3_3` to a **list** of all statements below that are **always** true.

1. An empirical histogram of the sample mean of a large random sample drawn with replacement from a population will be roughly normal.
2. If we know the mean and SD of any distribution, we can calculate a 68% confidence interval by stepping out one standard deviation from the mean in both directions.
3. For any distribution, 95% of the data falls within two standard deviations of the mean.
4. For any distribution, at least 95% of the data falls within five standard deviations of the mean.

```python
q3_3 = ...
```

```python
grader.check("q3_3")
```

## 4. Submission


**Important submission steps:** 
1. Run the tests and verify that they all pass.
2. Choose **Save Notebook** from the **File** menu, then **run the final cell**. 
3. Click the link to download the zip file.
4. Then submit the zip file to the corresponding assignment according to your instructor's directions. 

**It is your responsibility to make sure your work is saved before running the last cell.**

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_2 (points: [1, 1, 2])
```
>>> type(age_mean) == np.float64
True
```
```
>>> type(age_sd) == np.float64
True
```
```
>>> round(age_mean, 3) == 27.228
True
```


Tests for q1_3 (points: [1, 1, 2, 2])
```
>>> type(height_mean) == np.float64
True
```
```
>>> type(height_sd) == np.float64
True
```
```
>>> round(height_mean, 3) == 64.049
True
```
```
>>> round(height_sd, 3) == 2.525
True
```


Tests for q2_1 (points: [4])
```
>>> np.random.seed(8)
>>> np.round(one_sample_mean(salaries, 'salary', 100), 3) == 76699.829
True
```


Tests for q2_2 (points: [4])
```
>>> np.random.seed(8)
>>> val = simulate_sample_mean(salaries, 'salary', 100, 10, 0)
>>> plots.close("all")
>>> np.isclose(val, 4288.67990)
True
```


Tests for q2_4 (points: [0])
```
>>> type(q2_4) == np.ndarray
True
```


Tests for q2_5 (points: [0])
```
>>> type(q2_5) == np.ndarray
True
```


Tests for q2_6 (points: [0])
```
>>> type(SD_of_sample_means) == int
True
```


Tests for q2_7 (points: [0])
```
>>> type(pop_vs_sample) == np.ndarray
True
```


Tests for q3_1 (points: [0, 8])
```
>>> q3_1 in [1, 2, 3, 4]
True
```
```
>>> q3_1 == 2
True
```


Tests for q3_2 (points: [0])
```
>>> q3_2 in [1, 2, 3, 4]
True
```


Tests for q3_3 (points: [0])
```
>>> type(q3_3) == np.ndarray
True
```
