```python
# Initialize Otter
import otter
grader = otter.Notebook("hw08.ipynb")
```

# Homework 8: Probability and Simulation

Please complete this notebook by filling in the cells provided. Before you begin, execute the previous cell to load the provided tests. Each time you start your server, you will need to execute this cell again.

**Helpful Resource:**
- [Python Reference](https://pages.mtu.edu/~lebrown/data1202-s24/reference/index.html): Cheat sheet of helpful array & table methods used in DATA 1202!


**Recommended Readings**: 
* [Randomness](https://www.inferentialthinking.com/chapters/09/Randomness.html)
* [Sampling and Empirical Distributions](https://www.inferentialthinking.com/chapters/10/Sampling_and_Empirical_Distributions.html)



For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. **Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook!** For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!


**Note: This homework has hidden tests on it. That means even though the tests may say 100% passed, it doesn't mean your final grade will be 100%. We will be running more tests for correctness once everyone turns in the homework.**


Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. 

You should start early so that you have time to get help if you're stuck.

```python
# Run this cell to set up the notebook, but please don't change it.

# These lines import the Numpy and Datascience modules.
import numpy as np
from datascience import *
import d8error

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

## 1. Roulette

A Nevada roulette wheel has 38 pockets and a small ball that rests on the wheel. When the wheel is spun, the ball comes to rest in one of the 38 pockets. That pocket is declared the winner. 

The pockets are labeled 0, 00, 1, 2, 3, 4, ... , 36. Pockets 0 and 00 are green, and the other pockets are alternately red and black. The table `wheel` is a representation of a Nevada roulette wheel. **Note that *both* columns consist of strings.** Below is an example of a roulette wheel!

<img src="roulette_wheel.jpeg" width="330px">

Run the cell below to load the `wheel` table.

```python
wheel = Table.read_table('roulette_wheel.csv', dtype=str)
wheel
```

### Betting on Red ###
If you bet on *red*, you are betting that the winning pocket will be red. This bet *pays 1 to 1*. That means if you place a one-dollar bet on red, then:

- If the winning pocket is red, you gain 1 dollar. That is, you get your original dollar back, plus one more dollar.
- if the winning pocket is not red, you lose your dollar. In other words, you gain -1 dollars.

Let's see if you can make money by betting on red at roulette.

**Question 1.** Define a function `dollar_bet_on_red` that takes the name of a color and returns your gain in dollars if that color had won and you had placed a one-dollar bet on red. Remember that the gain can be negative. Make sure your function returns an integer. **(5 points)**

*Note:* You can assume that the only colors that will be passed as arguments are red, black, and green (as a string). Your function doesn't have to check that.


```python
def dollar_bet_on_red(...):
    ...
```

```python
grader.check("q1_1")
```

Run the cell below to make sure your function is working.

```python
print(dollar_bet_on_red('green'))
print(dollar_bet_on_red('black'))
print(dollar_bet_on_red('red'))
```

**Question 2.** Add a column labeled `Winnings: Red` to the table `wheel`. For each pocket, the column should contain your gain in dollars if that pocket won and you had bet one dollar on red. Your code should use the function `dollar_bet_on_red`. **(5 points)**


```python
red_winnings = ...
wheel = ...
wheel
```

```python
grader.check("q1_2")
```

### Simulating 10 Bets on Red
Roulette wheels are set up so that each time they are spun, the winning pocket is equally likely to be any of the 38 pockets regardless of the results of all other spins. Let's see what would happen if we decided to bet one dollar on red each round.

**Question 3.** Create a table `ten_bets` by sampling the table `wheel` to simulate 10 spins of the roulette wheel. Your table should have the same three column labels as in `wheel`. Once you've created that table, set `sum_bets` to your net gain in all 10 bets, assuming that you bet one dollar on red each time. **(5 points)**

*Hint:* It may be helpful to print out `ten_bets` after you create it!


```python
ten_bets = ...
sum_bets = ...(ten_bets.column('Winnings: Red'))
sum_bets
```

```python
grader.check("q1_3")
```

Run the cells above a few times to see how much money you would make if you made 10 one-dollar bets on red. Making a negative amount of money doesn't feel good, but it is a reality in gambling. Casinos are a business, and they make money when gamblers lose.

**Question 4.** Let's see what would happen if you made more bets. Define a function `net_gain_red` that takes the number of bets and returns the net gain in that number of one-dollar bets on red. **(5 points)**

*Hint:* You should use your `wheel` table within your function.


```python
def net_gain_red(...):
    ...
```

```python
grader.check("q1_4")
```

Run the cell below a few times to make sure that the results are similar to those you observed in the previous exercise.

```python
net_gain_red(10)
```

**Question 5.** Complete the cell below to simulate the net gain in 200 one-dollar bets on red, repeating the process 10,000 times. After the cell is run, `all_gains_red` should be an array with 10,000 entries, each of which is the net gain in 200 one-dollar bets on red. **(5 points)**

*Hint:* Think about which computational tool might be helpful for simulating a process multiple times. Lab 5 might be a good resource to look at!

*Note:* This cell might take a few seconds to run.


```python
num_bets = ...
repetitions = ...

all_gains_red = ...
...

len(all_gains_red) # Do not change this line! Check that all_gains_red is length 10000.
```

```python
grader.check("q1_5")
```

Run the cell below to visualize the results of your simulation.

```python
gains = Table().with_columns('Net Gain on Red', all_gains_red)
gains.hist(bins = np.arange(-80, 41, 4))
```

**Question 6:** Using the histogram above, decide whether the following statement is true or false:

>If you make 200 one-dollar bets on red, your chance of losing money is more than 50%.

Assign `loss_more_than_50` to either `True` or `False` depending on your answer to the question. **(5 points)**


```python
loss_more_than_50 = ...
```

```python
grader.check("q1_6")
```

### Betting on a Split ###
If betting on red doesn't seem like a good idea, maybe a gambler might want to try a different bet. A bet on a *split* is a bet on two consecutive numbers such as 5 and 6. This bets pays 17 to 1. That means if you place a one-dollar bet on the split 5 and 6, then:

- If the winning pocket is either 5 or 6, your gain is 17 dollars.
- If any other pocket wins, you lose your dollar, so your gain is -1 dollars.

**Question 7.** Define a function `dollar_bet_on_split` that takes a pocket number and returns your gain in dollars if that pocket won and you had bet one dollar on the 5-6 split. **(5 points)**

*Hint:* Remember that the pockets are represented as strings.


```python
def dollar_bet_on_split(...):
    ...
```

```python
grader.check("q1_7")
```

Run the cell below to check that your function is doing what it should.

```python
print(dollar_bet_on_split('5'))
print(dollar_bet_on_split('6'))
print(dollar_bet_on_split('00'))
print(dollar_bet_on_split('23'))
```

**Question 8.** Add a column `Winnings: Split` to the `wheel` table. For each pocket, the column should contain your gain in dollars if that pocket won and you had bet one dollar on the 5-6 split. **(5 points)**


```python
split_winnings = ...
wheel = ...
wheel.show(5) # Do not change this line.
```

```python
grader.check("q1_8")
```

**Question 9.** Simulate the net gain in 200 one-dollar bets on the 5-6 split, repeating the process 10,000 times and saving your gains in the array `all_gains_split`. **(5 points)**

*Hint:* Your code in Questions 4 and 5 may be helpful here!  For example, create a function `net_gain_split`, like you had before for `net_gain_red`.


```python
all_gains_split = ...

...

# Do not change the two lines below
gains = gains.with_columns('Net Gain on Split', all_gains_split)
gains.hist(bins = np.arange(-200, 150, 20))
```

```python
grader.check("q1_9")
```

**Question 10.** Look carefully at the visualization above, and assign `histogram_statements` to an array of the numbers of each statement below that can be correctly inferred from the overlaid histogram. **(5 points)**

1. If you bet one dollar 200 times on a split, your chance of losing money is more than 50%.
2. If you bet one dollar 200 times in roulette, your chance of making more than 50 dollars is greater if you bet on a split each time than if you bet on red each time.
3. If you bet one dollar 200 times in roulette, your chance of losing more than 50 dollars is greater if you bet on a split each time than if you bet on red each time.

*Hint:* We've already seen one of these statements in a prior question.


```python
histogram_statements = ...
```

```python
grader.check("q1_10")
```

If this exercise has put you off playing roulette, it has done its job. If you are still curious about other bets, [here](https://en.wikipedia.org/wiki/Roulette#Bet_odds_table) they all are, and [here](https://en.wikipedia.org/wiki/Roulette#House_edge) is the bad news. The house – that is, the casino – always has an edge over the gambler.

## 2. Chances

Before you do this exercise, make sure you understand the logic behind all the examples in [Section 9.5](https://inferentialthinking.com/chapters/09/5/Finding_Probabilities.html). 

Good ways to approach probability calculations include:

- Thinking one trial at a time: What does the first one have to be? Then what does the next one have to be?
- Breaking up the event into distinct ways in which it can happen.
- Seeing if it is easier to find the chance that the event does not happen.

### Finding Chances

On each spin of a roulette wheel, all 38 pockets are equally likely to be the winner regardless of the results of other spins. Among the 38 pockets, 18 are red, 18 black, and 2 green. In each part below, write an expression that evaluates to the chance of the event described.

**Question 1.** The winning pocket is black on all of the first three spins. **(5 points)**


```python
first_three_black = ...
```

```python
grader.check("q2_1")
```

**Question 2.** The color green never wins in the first 10 spins. **(5 points)**


```python
no_green = ...
```

```python
grader.check("q2_2")
```

**Question 3.** The color green wins at least once on the first 10 spins. **(5 points)**


```python
at_least_one_green = ...
```

```python
grader.check("q2_3")
```

**Question 4.** Two of the three colors never win in the first 10 spins. **(5 points)**

*Hint:* Imagine the event with the colors red, yellow, and orange. What situation(s) lead to two of the three colors never winning in the first 10 spins?


```python
lone_winners = ...
```

```python
grader.check("q2_4")
```

### Comparing Chances
In each of Questions 5-7, two events A and B are described. Choose from one of the following three options and set each answer variable to a single integer:

1. Event A is more likely than Event B
2. Event B is more likely than Event A
3. The two events have the same chance.

You should be able to make the choices **without calculation**. Good ways to approach this exercise include imagining carrying out the chance experiments yourself, one trial at a time, and by thinking about the [law of averages](https://inferentialthinking.com/chapters/10/1/Empirical_Distributions.html#the-law-of-averages).

**Question 5.** A child picks four times at random from a box that has four toy animals: a bear, an elephant, a giraffe, and a kangaroo. **(5 points)**

- Event A: all four different animals are picked (assuming the child picks without replacement)
- Event B: all four different animals are picked (assuming the child picks with replacement)


```python
toys_option = ...
```

```python
grader.check("q2_5")
```

**Question 6.** In a lottery, two numbers are drawn at random with replacement from the integers 1 through 1000. **(5 points)**

- Event A: The number 8 is picked on both draws
- Event B: The same number is picked on both draws


```python
lottery_option = ...
```

```python
grader.check("q2_6")
```

**Question 7.** A fair coin is tossed repeatedly. **(5 points)**

- Event A: There are 60 or more heads in 100 tosses
- Event B: There are 600 or more heads in 1000 tosses

*Hint*: Think about the law of averages!


```python
coin_option = ...
```

```python
grader.check("q2_7")
```

**Important submission steps:** 
1. Run the tests and verify that they all pass.
2. Choose **Save Notebook** from the **File** menu, then **run the final cell**. 
3. Click the link to download the zip file.
4. Then submit the zip file to the corresponding assignment according to your instructor's directions. 

**It is your responsibility to make sure your work is saved before running the last cell.**

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: [2])
```
>>> dollar_bet_on_red('red') == 1
True
```


Tests for q1_10 (points: [0, 0])
```
>>> # Make sure histogram_statements is an array.
>>> import numpy as np
>>> type(histogram_statements) == np.ndarray
True
```
```
>>> # Make sure you have set histogram_statements to an array with at least 1 number
>>> 1 <= histogram_statements.item(0) <= 3
True
```


Tests for q1_2 (points: [0])
```
>>> # Make sure your column names are correct
>>> wheel.labels[2] == "Winnings: Red"
True
```


Tests for q1_3 (points: [0])
```
>>> # If this test isn't passing, try running the cells from the top.
>>> set(["Pocket", "Color", "Winnings: Red"]) == set(ten_bets.labels)
True
```


Tests for q1_4 (points: [0])
```
>>> net_gain_red(10000) != net_gain_red(10000)
True
```


Tests for q1_5 (points: [0])
```
>>> len(all_gains_red) == 10000
True
```


Tests for q1_6 (points: [0])
```
>>> type(loss_more_than_50) == bool
True
```


Tests for q1_7 (points: [2])
```
>>> dollar_bet_on_split('5') == 17
True
```


Tests for q1_8 (points: [0])
```
>>> # Make sure your column names are correct
>>> wheel.labels[3] == "Winnings: Split"
True
```


Tests for q1_9 (points: [0])
```
>>> len(all_gains_split) == 10000
True
```


Tests for q2_1 (points: [0])
```
>>> # Your probability output should be a value between 0 and 1.
>>> 0 < first_three_black < 1
True
```


Tests for q2_2 (points: [0])
```
>>> # Your probability output should be a value between 0 and 1.
>>> 0 < no_green < 1
True
```


Tests for q2_3 (points: [0])
```
>>> # Your probability output should be a value between 0 and 1.
>>> 0 < at_least_one_green < 1
True
```


Tests for q2_4 (points: [0])
```
>>> # Your probability output should be a value between 0 and 1.
>>> 0 < lone_winners < 1
True
```


Tests for q2_5 (points: [0, 0])
```
>>> # Make sure you are setting toys_option to an int
>>> type(toys_option) == int
True
```
```
>>> # toys_option should be assigned to either 1, 2 or 3.
>>> 1 <= toys_option <= 3
True
```


Tests for q2_6 (points: [0, 0])
```
>>> # Make sure you are setting lottery_option to an int
>>> type(lottery_option) == int
True
```
```
>>> # toys_option should be assigned to either 1, 2 or 3.
>>> 1 <= lottery_option <= 3
True
```


Tests for q2_7 (points: [0, 0])
```
>>> # Make sure you are setting coin_option to an int
>>> type(coin_option) == int
True
```
```
>>> 1 <= coin_option <= 3
True
```
